```python
version = "REPLACE_PACKAGE_VERSION"
```

# Assignment 1 Part 2: Introduction to Supervised Machine Learning
In this assignment you will be using the [Breast Cancer Wisconsin (Diagnostic)](https://archive.ics.uci.edu/ml/datasets/Breast+Cancer+Wisconsin+(Diagnostic)) dataset to create a classifier that can help diagnose patients. We chose this data not only because it provides a good basis for a kNN classification problem, but also because it illustrates one of the built-in datasets that comes with `scikit-learn`.

```python
# Suppress all warnings
import warnings
warnings.filterwarnings('ignore')

# Import some necessary libararies 
import numpy as np
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
%matplotlib inline
```

```python
from sklearn.datasets import load_breast_cancer

# Load the dataset from scikit-learn. To see a description of the dataset, uncomment the print statement
cancer = load_breast_cancer()
# print(cancer.DESCR)
```

The object returned by `load_breast_cancer()` is a scikit-learn `Bunch` object, which is similar to a dictionary.

```python
# We can take a look at what essential attributes it has. Feel free to explore this object yourself. 
cancer.keys()
```

### Question 0. Warm-up (5 pts)

Understanding how many features you're dealing with and what each feature represents is an essential first step in machine learning.  So, how many features are there in this dataset? Complete the function below to return the answer as an integer. 

<!-- grade_id: cell-0562c45651419b03, points: 0, grade: false, solution: true -->
```python
def answer_zero():
    """
    This function returns the number of features of the breast cancer dataset as an integer. 
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-919dd7f0bfcbfd9a, points: 5, grade: true, solution: false -->
```python
# Hidden Tests Below

```

### Question 1. Data Transformation (10 pts)

In a lot of cases, raw data may not come in a form that's amenable to further manipulation or interpretation. Therefore, we may need to transform the raw data so that it better fits our purposes. In this assignment, we will store the data in a more human-friendly tabular format as a `pd.DataFrame`.

Complete the function below to return a `pd.DataFrame` of the shape `(569, 31)` with the following columns: 
```
['mean radius', 'mean texture', 'mean perimeter', 'mean area',
'mean smoothness', 'mean compactness', 'mean concavity',
'mean concave points', 'mean symmetry', 'mean fractal dimension',
'radius error', 'texture error', 'perimeter error', 'area error',
'smoothness error', 'compactness error', 'concavity error',
'concave points error', 'symmetry error', 'fractal dimension error',
'worst radius', 'worst texture', 'worst perimeter', 'worst area',
'worst smoothness', 'worst compactness', 'worst concavity',
'worst concave points', 'worst symmetry', 'worst fractal dimension',
'target']
```

    

<!-- grade_id: cell-7da0be2f62cc82c6, points: 0, grade: false, solution: true -->
```python
def answer_one():
    cancer_df = None
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return cancer_df
```

<!-- grade_id: cell-50152fa70665c294, points: 10, grade: true, solution: false -->
```python
# Autograder tests

stu_ans = answer_one()

assert stu_ans.shape == (569, 31), "Q1: The shape of your dataframe isn't correct. "

assert list(stu_ans.columns) == ['mean radius', 'mean texture', 'mean perimeter', 'mean area',
       'mean smoothness', 'mean compactness', 'mean concavity',
       'mean concave points', 'mean symmetry', 'mean fractal dimension',
       'radius error', 'texture error', 'perimeter error', 'area error',
       'smoothness error', 'compactness error', 'concavity error',
       'concave points error', 'symmetry error', 'fractal dimension error',
       'worst radius', 'worst texture', 'worst perimeter', 'worst area',
       'worst smoothness', 'worst compactness', 'worst concavity',
       'worst concave points', 'worst symmetry', 'worst fractal dimension',
       'target'], "Q1: Please check the column names of your dataframe."

del stu_ans
```

### Question 2. Class Distribution (5 pts) 

It's often a good idea to get some descriptive statistics, such as mean and variance of certain features, on the data at hand to understand the big picture. 

In particular, it's always a good idea to ask: what is the class distribution? That is, how many instances belong to the *malignant* class (encoded as 0) and the *benign* class (encoded as 1), respectively? Complete the function below to return the class distribution as a `pd.Series` of length 2 whose index is  `['malignant', 'benign']`

<!-- grade_id: cell-fac24fc5a519ef06, points: 0, grade: false, solution: true -->
```python
def answer_two():
    dist = None
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return dist
```

<!-- grade_id: cell-1ebcc2fe5b95b5df, points: 5, grade: true, solution: false -->
```python
# Autograder tests

stu_ans = answer_two()

assert isinstance(stu_ans, pd.Series), "Q2: Your result should be a pd.Series."


del stu_ans
```

### Question 3. Data Preparation (5 pts)

Training a classifier is a *supervised* machine learning problem, in which each instance $x_i$ has a corresponding class label $y_i$. All the instances $x_i$'s are collected into a matrix $X$ (with one instance per row of $X$), and all the corresponding labels are put into a column vector $y$. 

Now let's prepare the data for use with `scikit-learn`. Complete the function below to split our DataFrame into `X` (the data) and `y` (the labels), and to return them as a `tuple`, where
* `X` is a `pd.DataFrame` of the shape `(569, 30)`
* `y` is a `pd.Series` of the shape `(569,)`.

<!-- grade_id: cell-61b34fdd39e99b03, points: 0, grade: false, solution: true -->
```python
def answer_three():
    X, y = None, None

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return X, y
```

<!-- grade_id: cell-0bd5392f7412ea87, points: 5, grade: true, solution: false -->
```python
# Autograder tests

stu_ans = answer_three()

assert isinstance(stu_ans, tuple), "Q3: You should return a tuple!"
assert isinstance(stu_ans[0], pd.DataFrame), "Q3: X should be a pd.DataFrame. "
assert isinstance(stu_ans[1], pd.Series), "Q3: y should be a pd.Series. "
assert stu_ans[0].shape == (569, 30), "Q3: Please check the shape of X."
assert stu_ans[1].shape == (569,), "Q3: Please check the shape of y."

del stu_ans
```

### Question 4. Train-test Split (10 pts)

For a typical machine learning problem, we'd need two separate datasets, one for training a model and the other for evaluating the trained model for its generalisability to unseen data. `scikit-learn` provides a very handy [`train_test_split`](https://scikit-learn.org/stable/modules/generated/sklearn.model_selection.train_test_split.html) function for this purpose.  

Now, complete the function below that uses the `train_test_split` function to split `X` and `y` into training and testing sets. Your function should return a `tuple` `(X_train, X_test, y_train, and y_test)` where


* `X_train` is of the shape `(426, 30)`
* `X_test` is of the shape `(143, 30)`
* `y_train` is of the shape `(426,)`
* `y_test` is of the shape `(143,)`

**IMPORTANT: Set the random number generator state to the number 42 by specifying `random_state=42` to ensure a deterministic result that matches that of the autograder.**   Why the number 42?  Please see: https://www.theguardian.com/books/2011/feb/03/douglas-adams-42-hitchhiker

(In later work, we'll actually use a slightly more sophisticated splitting scheme that uses training, validation, and test sets, but we'll cover this later as part of a technique called cross-validation.)

<!-- grade_id: cell-8e5c85428bc3e008, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split

def answer_four():
    X_train, X_test, y_train, y_test = (None, ) * 4
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return X_train, X_test, y_train, y_test
```

<!-- grade_id: cell-7c3f4473c01b46df, points: 10, grade: true, solution: false -->
```python
# Autograder tests

stu_ans = answer_four()

assert stu_ans[0].shape == (426, 30), "Q4: Please check the shape of X_train."
assert stu_ans[1].shape == (143, 30), "Q4: Please check the shape of X_test."
assert stu_ans[2].shape == (426,), "Q4: Please check the shape of y_train."
assert stu_ans[3].shape == (143,), "Q4: Please check the shape of y_test."

del stu_ans
```

### Question 5. kNN Classifier Construction (5 pts)

Use `KNeighborsClassifier` from `scikit-learn` to fit a $k$-Nearest Neighbours ($k$NN) classifier with `X_train` and `y_train` where $k = 1$. Your function should return the trained classifier itself, which is a `sklearn.neighbors.KNeighborsClassifier` object.

<!-- grade_id: cell-4b71f9b9bd2e9139, points: 0, grade: false, solution: true -->
```python
from sklearn.neighbors import KNeighborsClassifier

def answer_five():
    knn = None
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return knn
```

<!-- grade_id: cell-746df875c076339d, points: 5, grade: true, solution: false -->
```python
# Autograder tests

stu_ans = answer_five()
assert isinstance(stu_ans, KNeighborsClassifier), "Q5: Please build the required kNN classifier."
assert len(stu_ans.classes_) == 2, "Q5: Your kNN classifier was trained with an incorrect # classes. "
del stu_ans
```

### Question 6. kNN Prediction on Mean Vector (10 pts)

It's often useful and interesting to know what class a "typical" or "average" data point belongs to. Use your kNN classifier from the last question to predict the class label for the *mean vector* of the training data. Your function should return the predicted class label as a singleton numpy array --- either `array([ 0.])` or `array([ 1.])`. 

If you encounter errors complaining that the shape of your data isn't correct, carefully check the [documentation](https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.KNeighborsClassifier.html#sklearn.neighbors.KNeighborsClassifier.predict) of the `predict` function. Another useful hint is to consider what kind of object `X_train` is. How do you make sure it is in the correct shape?

<!-- grade_id: cell-bbc431b30a17745f, points: 0, grade: false, solution: true -->
```python
def answer_six():
    pred = None

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return pred
```

<!-- grade_id: cell-c93e4911f581b5cb, points: 10, grade: true, solution: false -->
```python
# Autograder tests

stu_ans = answer_six()

assert isinstance(stu_ans, np.ndarray), "Q6: Your function should return a np.ndarray. "

del stu_ans
```

### Question 7. kNN Prediction on the Test Set (5 pts)

Now, use your kNN classifier to predict class labels for the test set `X_test`. Your function should return a binary `np.ndarray` of the shape `(143,)` whose values are either `0.0` or `1.0`

<!-- grade_id: cell-556e3b89a523853a, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    preds = None
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return preds
```

<!-- grade_id: cell-1fc7ffa2d7e05814, points: 5, grade: true, solution: false -->
```python
# Autograder tests

stu_ans = answer_seven()
assert isinstance(stu_ans, np.ndarray), "Q7: Your function should return a np.ndarray. "

del stu_ans
```

### Question 8. Evaluation on the Test Set (5 pts) 

Once you have the predictions on the test set, you may compare them with the ground-truth labels to gauge how well your model performs when given unseen data. 

Complete the function below to compute the score (mean accuracy) of your kNN classifier using the test set `X_test` and the test labels `y_test`. The function should return a `float` between 0 and 1.

<!-- grade_id: cell-dc56668e96e99a63, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    score = None
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return score
```

<!-- grade_id: cell-edfa0ff11db4a6d9, points: 5, grade: true, solution: false -->
```python
# Autograder tests

stu_ans = answer_eight()

assert isinstance(stu_ans, float), "Q8: Your function should return a float. "

del stu_ans
```

### Optional plot

Note: The following plots will help you to evaluate the model with showing the prediction accuracy in training and testing set.

Try using the plotting function below to visualize the different predicition scores between training and test sets, as well as malignant and benign cells.

```python
def accuracy_plot(knn):
#     import matplotlib.pyplot as plt

#     %matplotlib notebook

    X_train, X_test, y_train, y_test = answer_four()

    # Find the training and testing accuracies by target value (i.e. malignant, benign)
    mal_train_X = X_train[y_train==0]
    mal_train_y = y_train[y_train==0]
    ben_train_X = X_train[y_train==1]
    ben_train_y = y_train[y_train==1]

    mal_test_X = X_test[y_test==0]
    mal_test_y = y_test[y_test==0]
    ben_test_X = X_test[y_test==1]
    ben_test_y = y_test[y_test==1]

    scores = [knn.score(mal_train_X, mal_train_y), knn.score(ben_train_X, ben_train_y), 
              knn.score(mal_test_X, mal_test_y), knn.score(ben_test_X, ben_test_y)]


    plt.figure()

    # Plot the scores as a bar chart
    bars = plt.bar(np.arange(4), scores, color=['#4c72b0','#4c72b0','#55a868','#55a868'])

    # directly label the score onto the bars
    for bar in bars:
        height = bar.get_height()
        plt.gca().text(bar.get_x() + bar.get_width()/2, height*.90, '{0:.{1}f}'.format(height, 2), 
                     ha='center', color='w', fontsize=11)

    # remove all the ticks (both axes), and tick labels on the Y axis
    plt.tick_params(top='off', bottom='off', left='off', right='off', labelleft='off', labelbottom='on')

    # remove the frame of the chart
    for spine in plt.gca().spines.values():
        spine.set_visible(False)

    plt.xticks([0,1,2,3], ['Malignant\nTraining', 'Benign\nTraining', 'Malignant\nTest', 'Benign\nTest'], alpha=0.8);
    plt.title('Training and Test Accuracies for Malignant and Benign Cells', alpha=0.8)
```

Uncomment the below line to see the visualization. You can pass in any trained classifier as an argument.

**Comment out** the the below line when submitting your notebook for grading. 

```python
# Remember to comment it out when submitting the notebook
# accuracy_plot(answer_five())
```

### Question 9. Hyper-parameter Tuning (5 pts)

It's common to experiment with different configurations of a model, also known as "hyper-parameters" which are typically specified in advance of starting the step of estimating the learnable "parameters" of a specific model configuration in order to achieve better performance. The main crucial hyper-parameter of a k-NN model is the number $k$, the number of neighbors to examine. 

Change $k$ to 15 and fit the model with training data. Complete the function below to return the trained model, which is a `sklearn.neighbors.KNeighborsClassifier` object. Compare this with the results above for $k$ = 1.

<!-- grade_id: cell-3bc80326a3e99e05, points: 0, grade: false, solution: true -->
```python
from sklearn.neighbors import KNeighborsClassifier

def answer_nine():
    knn = None
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return knn

```

<!-- grade_id: cell-e724a46e381ed1b7, points: 5, grade: true, solution: false -->
```python
# Autograder tests

stu_ans = answer_nine()
assert isinstance(stu_ans, KNeighborsClassifier), "Q9: Please build the required kNN classifier."
assert len(stu_ans.classes_) == 2, "Q9: Your kNN classifier was trained with an incorrect # classes. "
del stu_ans
```

```python
# Remember to comment it out when submitting the notebook
# accuracy_plot(answer_nine())
```

### Question 10. Weighted k-NN (5 pts)

Keeping $k = 15$, now change the k-NN method to use a *weighted* distance measure: this means closer neighbors of a query point will have more influence on the prediction than neighbors which are a greater distance away. (Normally, the default k-NN classifier uses a uniform weighting, i.e. it ignores how far a neighbor is and just sees that it exists.)

Your function below should return a trained kNN classifier of the type `sklearn.neighbors.KNeighborsClassifier`. (You may find it helpful to plot and compare the results with the unweighted distance measure, using the provided plotting function.)

<!-- grade_id: cell-4243c4d7cfa2960e, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    knn = None

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return knn
```

<!-- grade_id: cell-558b71c1a69189f5, points: 5, grade: true, solution: false -->
```python
# Autograder tests

stu_ans = answer_ten()
assert isinstance(stu_ans, KNeighborsClassifier), "Q9: Please build the required kNN classifier."
assert len(stu_ans.classes_) == 2, "Q9: Your kNN classifier was trained with an incorrect # classes. "
del stu_ans
```

```python
# Remember to comment it out when submitting the notebook
# accuracy_plot(answer_ten())
```

### Question 11: Model parameter tuning (15 points)

It is important to improve algorithm design through automatically changing the parameter based on data-driven evidence, because it is more effective than just 'trying out' different parameters by hand.

Perform a simple parameter sweep for all **odd** values of $k$ from 1 to 19 inclusive, and return the optimal value of $k$ that leads to the highest overall *test set accuracy* on this train/test split.  Accuracy is computed using the **score** method. Your code should return an integer between 1 and 19. In case of a tie, return the smallest best $k$.

<!-- grade_id: cell-feeb5e4810d6c9e6, points: 0, grade: false, solution: true -->
```python
def answer_eleven():
    k_best = None

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return k_best
```

<!-- grade_id: cell-e17d6199583d1080, points: 15, grade: true, solution: false -->
```python
# Autograder tests

stu_ans = answer_eleven()

assert isinstance(stu_ans, int), "Q11: Your function should return an integer. "

del stu_ans
```

### Question 12: Overfitting  (15 points)

A key sign of overfitting is obtaining a training set accuracy that is extremely high (or even perfect), but a test set accuracy that is considerably lower. It is tempting to use a classifier that does so well on the training data it was given, but reality sets in when we try to use this overfit classifier on new test data and we discover it does not in fact generalize well.

Using the same set of possible k-NN classifiers and values of $k$ as the previous question (Q11), look for a scenario where overfitting is likely to be happening, by finding the optimal value for $k$ if your goal was to pick the classifier that did best only on the **training set**.  Compute what the resulting test set accuracy would have been, if you had picked that training-set-based value for $k$.

Your function should return an (`int`, `float`, `float`) tuple, as follows:

`tuple[0]`: the optimal value of $k$ that maximizes *training set* accuracy

`tuple[1]`: the corresponding *training set* accuracy for that optimal $k$

`tuple[2]`: the corresponding *test set* accuracy that you would have received *if* you had used that optimal $k$.

(It is instructive to compare this test set accuracy with the best one you were able to achieve in question 11.)

<!-- grade_id: cell-fa4b5a5c9ec7dd40, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    k_best = None

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return k_best
```

<!-- grade_id: cell-9d8ba54a5e282cd7, points: 15, grade: true, solution: false -->
```python
# Autograder tests

stu_ans = answer_twelve()

assert isinstance(stu_ans, tuple), "Q12: Your function should return a tuple. "
assert len(stu_ans) == 3, "Q12: The length of your returned tuple should be 3. "
assert isinstance(stu_ans[0], int), "Q12: Your tuple format should be (*int*, float, float). "
assert isinstance(stu_ans[1], float), "Q12: Your tuple format should be (int, *float*, float). "
assert isinstance(stu_ans[1], float), "Q12: Your tuple format should be (int, float, *float*). "


del stu_ans
```